---
title: "Tidy Data"
---

### Introduction
As IO Psychologists, we often deal with a lot of data, be it employee performance metrics, survey responses, or recruitment data. One of the key challenges we face is ensuring that this data is well-organized and easy to analyze. Enter the concept of "Tidy Data" - a term coined by Hadley Wickham, the creator of many of the tools in the tidyverse package in R.

According to Wickham, "Tidy datasets are all alike, but every messy dataset is messy in its own way." The idea is to organize your data in a consistent way that makes it easier to work with and analyze. In essence, tidy data is about structuring datasets to facilitate analysis.  

### Tidy Data Principles

Tidy data follows three fundamental rules:  

- Each variable must have its own column.  
- Each observation must have its own row.  
- Each value must have its own cell.  


<figure>
  <img src="images/tidy-principles.png" width="100%">
  <figcaption class="figcap">Tidy Data Principles</figcaption>
</figure>


###
The advantages of tidy data are twofold. Firstly, having a consistent data structure makes it easier to learn the tools that work with it. Secondly, placing variables in columns allows R's vectorized nature to shine, making data transformation feel natural.\

By maintaining tidy datasets, you can streamline your data analysis process, make it easier to handle missing values, and enhance the readability and organization of your data.  


###
Let's consider an example where we have data on employee satisfaction, performance, and tenure. The tidy version of this data would look something like this:  

| Employee ID | Satisfaction | Performance | Tenure |
|---|---|---|---|
| 1 | 7 | 8 | 2 |
| 2 | 8 | 7 | 3 |
| 3 | 9 | 9 | 5 |
| 4 | 6 | 5 | 1 | 


###
In this example, each variable (Satisfaction, Performance, Tenure) has its own column, each observation (each employee) has its own row, and each value has its own cell.  


### Pivoting Data

In practice, you'll often encounter data that is not tidy. For instance, a variable might be spread across multiple columns, or one observation might be scattered across multiple rows. To handle such scenarios, `tidyr`, a core tidyverse package for working with messy data, provides two pivotal functions: `pivot_longer()` and `pivot_wider()`.  

Let's discuss each one in detail.  

### 
**`pivot_longer()`**\ 

`pivot_longer()` is used to make "wide" data longer. In other words, it takes multiple columns and combines them into key-value pairs, duplicating all other columns as needed.  

```{r, eval=FALSE}
tidy_survey_data <- survey_data |> pivot_longer(
  cols = starts_with("Question"),
  names_to = "Question",
  values_to = "Answer"
)
```

###
**`pivot_wider()`**\

`pivot_wider()` does the opposite. It makes "long" data wider, spreading out a key-value pair across multiple columns.  

```{r, eval=FALSE}
wide_survey_data <- tidy_survey_data |> pivot_wider(
  names_from = "Question",
  values_from = "Answer"
)
```


### Separate and unite
The `separate()` function from the `tidyr` package is used to split a column into two columns at the point of a specified separator also called delimiter (e.g., comma, space, semi-colon). This is useful when for breaking down combined information into distinct parts, which can be valuable when analyzing employee data.


###
In the following example, we have an `employee_data_sep` dataset that includes a column for full names. First, type the variable name in the codeblock and view the dataset:\

```{r prepare-separate-data}
employee_data_sep <- tibble(
  employee_id = c(1, 2, 3),
  full_name = c("John Doe", "Jane Smith", "Michael Johnson")
)
```

```{r view-separate-data, exercise = TRUE, exercise.setup = "prepare-separate-data"}

```

```{r view-separate-data-solution}
employee_data_sep
```

###
Now, complete the code in the codeblock below so that it calls `separate()` and splits the `full_name` column into separate columns for `first_name` and `last_name`. Run the code to inspect your result and then click **Submit Answer**.\

*Remember to run `?help` in your console if you need to read the `separate()` documentation*.\

```{r separate, exercise = TRUE, exercise.setup = "prepare-separate-data"}
employee_data_sep |> 
  ___(
    ___, 
    into = c("___", "___"), 
    sep = ___
  )
```

```{r separate-hint-1}
employee_data_sep |> 
  separate(
    full_name, 
    into = c("first_name", "___"), 
    sep = " "
  )
```

```{r separate-solution, exercise.reveal_solution = FALSE}
employee_data_sep |> 
  separate(
    full_name, 
    into = c("first_name", "last_name"), 
    sep = " "
  )
```

```{r separate-check}
grade_this({
    pass_if_equal()
    fail()
})
```


###
Conversely, we might need to combine two columns into one. The `unite()` function does this. We can save the `employee_data_sep` data from the previous example in a new variable called `employee_data_unite` and then call `unite()` to combine the `first_name` and `last_name` columns back into a single `full_name` column. Execute this task by completing the code in the codeblock below. Run the code to inspect your result, and then click **Submit Answer** to see if your answer is correct.\

*Remember to run `?help` in your console if you need to read the `unite()` documentation*.\

```{r prepare-unite-data}
employee_data_sep <- tibble(
    employee_id = c(1, 2, 3),
    full_name = c("John Doe", "Jane Smith", "Michael Johnson")
  ) |> 
    separate(
      full_name, 
      into = c("first_name", "last_name"), 
      sep = " "
    )
```

```{r unite, exercise = TRUE, exercise.setup = "prepare-unite-data"}
employee_data_unite <- employee_data_sep |> 
  ___(
    col = "___",
    c(___, ___),
    sep = ___
  )

employee_data_unite
```

```{r unite-hint-1}
employee_data_unite <- employee_data_sep |> 
  ___(
    col = "full_name",
    c(first_name, last_name),
    sep = " "
  )

employee_data_unite
```

```{r unite-solution, exercise.reveal_solution = FALSE}
employee_data_unite <- employee_data_sep |> 
  unite(
    col = "full_name",
    c(first_name, last_name),
    sep = " "
  )

employee_data_unite
```

###
Great work!\
For the next and final section of this lesson, we will dive a bit deeper into one of the most useful and invaluable methods for working with the `tidyverse`: The pipe operator (`|>`).\


### The Pipe Operator **(`|>`)**  
I bet you are wondering what that strange `|>` symbol is that we used in the above section. This symbol is a very useful method for creating tidy data, and it called the pipe operator (`|>`). This operator allows you to pipe your data from one function to the next. It's used to make your code cleaner and easier to read. The pipe operator fundamentally transforms the way you use R for data manipulation, leading to a tidier, more readable, and maintainable code.  

###
**Don't Panic!**\

You may be unfamiliar with some of the functions used in the following code examples (such as `filter()`). Don't worry! We will cover them in detail in the next lesson: Transforming Data!  

###
Consider an example where you have a dataset of employees, and you want to filter out the HR department, select the age and salary columns, and calculate the average salary. We can use a pipeline using `|>` that clearly and understandibly describes each step:  


```{r, eval=FALSE}
employees_data |> # step 1: dataset of employees
  filter(department == "HR") |> # step 2: filter out the HR department
  select(age, salary) |> # step 3: select the age and salary columns
  summarise(avg_salary = mean(salary, na.rm = TRUE)) # step 4: calculate the average salary
```

###
Notice how easy it is to understand what the code is executing. Each step is written on its own line. If you need to edit or change the code, this format simplifies the process.\

### The Benefit of Using Piping in Tidyverse

The introduction of the pipe operator has transformed the way we can write and read R code. Let's consider the same task as above, but without using the pipe operator:  

```{r, eval=FALSE}
avg_salary <- summarise(select(filter(employees_data, department == "HR"), age, salary), avg_salary = mean(salary, na.rm = TRUE))
```

###
Without piping, the above code becomes much harder to read. Functions are nested within one another, and the code reads from inside-out, rather than from left to right. It's also much more challenging to add or remove steps from the analysis.  

###
On the other hand, the pipe operator allows us to follow the logic of our data transformation step by step. It facilitates a style of programming that emphasizes clear, readable code, and promotes the chaining of functions, where the output of one function is the input to the next. Take another look at the piping of the `employees_data` to calculate the same `avg_salary`:  

```{r, eval=FALSE}
employees_data |> 
  filter(department == "HR") |> 
  select(age, salary) |> 
  summarise(avg_salary = mean(salary, na.rm = TRUE))
```


###
So much neater, right!? This is a powerful style, especially for data manipulation tasks.


### Conclusion
**Congratulations on completing this lesson!**\

You've taken a significant step in understanding the principles of tidy data. Test out your newfound knowledge by completing the "Tidy Data: Exercises" sub-lesson!\

The next lesson, titled "Transforming Data," will unlock new avenues by covering data transformation techniques. Prepare to embark on this exciting journey!
